Are peer reviewers influenced by their work being cited?

Author response:

The following is the authors’ response to the previous reviews.

Editors comments:

I would encourage you to submit a revised version that addresses the following two points:

[a] The point from Reviewer #1 about a possible major confounding factor. The following article might be germane here: Baas and Fennell, 2019: https://papers.ssrn.com/sol3/papers.cfm?abstract_id=3339568

I don’t believe that the point raised by reviewer 1 is a confounder, see my response below.

This article highlighted was in my reading list, but I did not cite it because I was confused by its methods.

The point from Reviewer #4 about the abstract. It is important that the abstract says something about how reviewers reacted to the original versions of articles in which they were cited (ie, the odds ratio = 0.84, etc result), before going on to discuss how they reacted to revised articles (ie, the odds ratio = 1.61, etc result). I would suggest doing this along the following lines - but please feel free to reword the passage "but this effect was not strong/conclusive":

When reviewers were cited in the original version of the article under review, they were less likely to approve the article compared with reviewers who were not cited, but this effect was not strong/conclusive (odds ratio = 0.84; adjusted 99.4% CI: 0.69-1.03). However, when reviewers were cited in the revised version of the article, they were more likely to approve compared with reviewers who were not cited (odds ratio = 1.61; adjusted 99.4% CI: 1.16-2.23).

I have changed the abstract to include the odds ratios for version 1 and have used the same wording as from the main text.

Reviewer #1 (Public review):

Summary:

The work used open peer reviews and followed them through a succession of reviews and author revisions. It assessed whether a reviewer had requested the author include additional citations and references to the reviewers' work. It then assessed whether the author had followed these suggestions and what the probability of acceptance was based on the authors decision. Reviewers who were cited were more likely to recommend the article for publication when compared with reviewers that were not cited. Reviewers who requested and received a citation were much likely to accept than reviewers that requested and did not receive a citation.

Strengths and weaknesses:

The work's strengths are the in-depth and thorough statistical analysis it contains and the very large dataset it uses. The methods are robust and reported in detail.

I am still concerned that there is a major confounding factor: if you ignore the reviewers requests for citations are you more likely to have ignored all their other suggestions too? This has now been mentioned briefly and slightly circuitously in the limitations section. I would still like this (I think) major limitation to be given more consideration and discussion, although I am happy that it cannot be addressed directly in the analysis.

This is likely to happen, but I do not think it’s a confounder. A confounder needs to be associated with both the outcome and the exposure of interest. If we consider forthright authors who are more likely to rebuff all suggestions, then they would receive just as many citation and self-citation requests as authors who were more compliant. The behaviour of forthright authors would likely only reduce the association seen in most authors which would be reflected in the odds ratios.

Reviewer #2 (Public review):

Summary:

This article examines reviewer coercion in the form of requesting citations to the reviewer's own work as a possible trade for acceptance and shows that, under certain conditions, this happens.

Strengths:

The methods are well done and the results support the conclusions that some reviewers "request" self-citations and may be making acceptance decisions based on whether an author fulfills that request.

Weakness:

I thank the author for addressing my comments about the original version.

Reviewer #3 (Public review):

Summary:

In this article, Barnett examines a pressing question regarding citing behavior of authors during the peer review process. In particular, the author studies the interaction between reviewers and authors, focusing on the odds of acceptance, and how this may be affected by whether or not the authors cited the reviewers' prior work, whether the reviewer requested such citations be added, and whether the authors complied/how that affected the reviewer decision-making.

Strengths:

The author uses a clever analytical design, examining four journals that use the same open peer review system, in which the identities of the authors and reviewers are both available and linkable to structured data. Categorical information about the approval is also available as structured data. This design allows a large scale investigation of this question.

Weaknesses:

My original concerns have been largely addressed. Much more detail is provided about the number of documents under consideration for each analysis, which clarifies a great deal.

Much of the observed reviewer behavior disappears or has much lower effect sizes depending on whether "Accept with Reservations" is considered an Accept or a Reject. This is acknowledged in the results text. Language has been toned down in the revised version.

The conditional analysis on the 441 reviews (lines 224-228) does support the revised interpretation as presented.

No additional concerns are noted.

Reviewer #4 (Public review):

Summary:

This work investigates whether a citation to a referee made by a paper is associated with a more positive evaluation by that referee for that paper. It provides evidence supporting this hypothesis. The work also investigates the role of self-citations by referees where the referee would ask authors to cite the referee's paper.

Strengths:

This is an important problem: referees for scientific papers must provide their impartial opinions rooted in core scientific principles. Any undue influence due to the role of citations breaks this requirement. This work studies the possible presence and extent of this.

The methods are solid and well done. The work uses a matched pair design which controls for article-level confounding and further investigates robustness to other potential confounds.

Weaknesses:

The authors have addressed most concerns in the initial review. The only remaining concern is the asymmetric reporting and highlighting of version 1 (null result) versus version 2 (rejecting null). For example the abstract says "We find that reviewers who were cited in the article under review were more likely to recommend approval, but only after the first version (odds ratio = 1.61; adjusted 99.4% CI: 1.16 to 2.23)" instead of a symmetric sentence "We find ... in version 1 and ... in version 2".

The latest version now includes the results for both versions.

Changes from the earlier version to the revised version

Abstract

Earlier version: Peer reviewers sometimes comment that their own journal articles should be cited by the journal article under review. Comments concerning relevant articles can be justified, but comments can also be unrelated coercive citations. Here, we used a matched study design to explore how citations influence the peer review process. We used a sample of more than 37,000 peer reviews from four journals that use open peer review and make all article versions available. We find that reviewers who were cited in the article under review were more likely to recommend approval, but only after the first version (odds ratio = 1.61; adjusted 99.4% CI: 1.16 to 2.23). For all versions of the articles, the reviewers who commented that their own articles should be cited were less likely to recommend approval compared to the reviewers who did not, with the strongest association after the first version (odds ratio = 0.15; adjusted 99.4% CI: 0.08 to 0.30). Reviewers who included a citation to their own articles were much more likely to approve a revised article that cited their articles compared to a revised article that did not (odds ratio = 3.5; 95% CI: 2.0 to 6.1). Some reviewers’ recommendations depend on whether they are cited or want to be cited. Reviewer citation requests can turn peer review into a transaction rather than an objective critique of the article.

Revised version: Peer reviewers sometimes comment that their own journal articles should be cited by the journal article under review. Comments concerning relevant articles can be justified, but comments can also be unrelated coercive citations. Here, we used a matched observational study design to explore how citations influence the peer review process. We used a sample of more than 37,000 peer reviews from four journals that use open peer review and make all article versions available. We find that reviewers who were cited in versions after version 1 were more likely to make a favourable recommendation (odds ratio = 1.61; adjusted 99.4% CI: 1.16 to 2.23), whereas being cited in the first version did not improve their recommendation (odds ratio = 0.84; adjusted 99.4% CI: 0.69 to 1.03). For all versions of the articles, the reviewers who commented that their own articles should be cited were less likely to recommend approval compared to the reviewers who did not, with the strongest association after the first version (odds ratio = 0.15; adjusted 99.4% CI: 0.08 to 0.30). Reviewers who included a citation to their own articles were much more likely to approve a revised article that cited their articles compared to a revised article that did not (odds ratio = 3.5; 95% CI: 2.0 to 6.1). Some reviewers’ recommendations depend on whether they are cited or want to be cited. Reviewer citation requests can turn peer review into a transaction rather than an objective critique of the article.

Introduction

In 2024, a published peer reviewed article included this remarkable sentence: “As strongly requested by the reviewers, here we cite some references [35–47] although they are completely irrelevant to the present work” [1]. This was a rare public example of coerced citations, where a reviewer exploits the peer review process to increase their citation counts and hence further their own career [2–4]. Reviewers should be relevant experts, so some suggestions to cite their articles will be appropriate. However, excessive citation requests or requests to cite unrelated articles are unethical [5–9]. Coerced citations can also come from editors trying to boost their journal’s ranking [10–12].

Coerced citations are reported as a common problem in peer review. In author surveys, two-thirds reported pressure from peer reviewers to cite unrelated articles [13] and 23% had experienced a reviewer that “required them to include unnecessary references to their publication(s)” [14]. Publishers have investigated whether “hundreds of researchers” have manipulated the peer review process to increase their own citations [15]. Some reviewers may be exploiting their power over authors who “have a strong incentive to […] accept all ‘suggestions’ by the referees even if one knows that they are misleading or even incorrect” [16].

As reviewers are often in the same field as the article’s authors, they may already be cited in the article without the need for coerced citations. Reviewers who are cited may give a more favourable peer review and be more willing to overlook flaws [17, 18]. Some authors may try to exploit this using “referee baiting” [3] or “flattery citations” [19] by favourably citing a reviewer’s work.

The interactions during peer review between authors and reviewers can determine whether an article is accepted [20] and what results are included in the published version [21]. Given the importance of peer review for science, studies that examine how peer review works in practice are needed [22–26]. Here, we examine interactions between peer reviewers and authors using four journals that publish all article versions and all peer reviews. We had two research questions:

Do peer reviewers give a more or less favourable recommendation when they are cited in the article?

Do peer reviewers give a more or less favourable recommendation when their review includes a citation to their own articles?

Methods

Journal selection

We studied journals from the publisher F1000 as their journals use open peer review with signed reviewers. F1000 journals use a publish-review-curate model [27], meaning all versions of the article are publicly available, including versions updated after peer review. This allowed us to examine the interactions between authors and reviewers throughout the peer review process. We selected four F1000 journals that each had over 100 articles. Some characteristics of the four journals are given in Table 1. Three journals were created to support funders.

The peer review process used by F1000 journals differs from most standard journals. The journals do not use academic editors, but do have in-house editors who manage articles but do not make editorial decisions. This means that most interactions during peer review are between authors and reviewers directly. In-house editors perform checks prior to the first version of the article being published and at F1000Research this results in 40 to 50% submissions being rejected (personal communication, F1000 staff). Up to mid-2024, authors were asked to identify potential reviewers who were qualified experts with no competing interests [28]. Since mid-2024, reviewer identification is made in-house, although authors can suggest reviewers.

Reviewers are asked to recommend one of three categories: Approved, Approved with reservations, and Not approved. For brevity, we refer to ‘Approved with Reservations’ as ‘Reservations’. An article is indexed once it receives two ‘Approved’ or two ‘Reservations’ and one ‘Approved’. The guidelines for recommending Approved are: “the aims and research methods are adequate; results are presented accurately, and the conclusions are justified and supported by the presented data” [29]. Peer reviewers are asked to assess the validity of an article’s content, rather than novelty or interest levels, an approach designed to combat publication bias [30].

All four journals have a peer reviewer code of conduct and state that reviewers should familiarise themselves with the ethical guidelines for peer reviewers by the Committee On Publication Ethics [31]. The journals’ guidelines for reviewers include the following: “reviewers should explicitly state their reasoning when asking authors to cite their own work”.

Data extraction

We extracted data on authors and articles from the OpenAlex database (https://openalex.org/) and directly from the four journals. OpenAlex combines scholarly data from multiple sources, including ORCID - a unique identifier for researchers, Microsoft Academic, Crossref and PubMed. A recent study compared OpenAlex with the two commonly used proprietary bibliometric databases of Web of Science and Scopus for the years 2015 to 2022 [32]. The results were mixed, but OpenAlex had better ORCID coverage and covered more Digital Object Identifiers (DOIs) - the unique identifier for publications. We accessed OpenAlex using the openalexR package [33, 34]. We used each journal’s application programming interface (API) to extract data on the articles and peer reviews. The data were extracted in four stages:

Searches were made using the APIs of the four journals to find all articles published between 1 Jan 2012 and 28 May 2025, with the start date to capture all potential articles.

For each article, the following data were downloaded in XML format:

The article’s publication date and version number

The reviewers’ names and ORCIDs (if available)

The text of all reviews and the reviewers’ recommendations

The DOIs and PMIDs (PubMed IDs) from the article’s reference list

The DOIs and PMIDs of any articles cited by the reviewers. The online peer review system at F1000 journals includes the DOI of any article cited in the review, which facilitates the identification of citations to the reviewers’ articles.

Articles were excluded if:

They were not peer reviewed or had yet to receive any reviews

The reference list was empty

The reviewers’ publication histories were collected from OpenAlex using their name, institution and ORCID (if available). Reviews were excluded if there was no record for the reviewer in OpenAlex, or if the reviewer had no published articles as there was no potential for them to be cited or request a citation to their own articles.

Study design

We used two predictor variables about the reviewer:

The number of times they were cited in the article (0,1, 2,…).

The number of times they included citations to their own articles in their review (0,1, 2,…).

We fitted both predictors as linear, but reviewers may behave differently with any citation rather than a linear change, and hence we also fitted both predictors as a binary “none versus any” (0 versus 1, 2,…). We compared the linear and binary alternatives using the Akaike Information Criterion (AIC) to find the parameterisation that best fitted the data [35].

We matched on article and version to control for confounding by any characteristics of the article [36]; for example, the article’s topic or writing style. Hence, we compared two or more independent reviewers who considered the same article.

All analyses were stratified by article version, using the first version only or the second and subsequent versions. This is because the reviewers are unknown to the authors for the first version, but from the second version onwards, the authors will know the reviewers as the journals use signed peer reviews. This knowledge could alter the behaviour of authors and reviewers.

The study design is summarised in Figure 1.

Graphical summary of the study design for research question 1 showing a dummy article and two reviews.

In the first version of the article, the reviewer Smith (blue) is cited whilst Jones (purple) is not. For the second version of the article, the authors are now aware that Jones is a reviewer and Jones has been cited. The reviewers’ recommendations are the outcome and are colour-coded as Not approved (red), Reservations (orange) and Approved (green). We tested whether citations to the reviewer in the article influenced their recommendation. The matched design means that only reviewers of the same article are compared (here, Smith and Jones) and the overall effect is estimated by aggregating over multiple matched comparisons. Research question 2 used the same design but examined citations to the reviewers’ articles in their reviews.

Statistical methods

We used conditional logistic regression to examine the associations between the citations to the reviewer and their ordinal recommendation (Approved → Reservations → Not approved) while matching the article and the version [37]. Conditional logistic regression requires a binary dependent variable; hence, we fitted two related models that examined the odds of:

“Approved” compared with “Reservations” or “Not approved”.

“Approved” or “Reservations” compared with “Not approved”,

These two models tested the same hypothesis, hence we adjusted for multiple testing. We also used repeated testing due to the stratification by article version and the two formulations of the predictors (linear or none versus any). Since we used 8 (2 × 2 × 2) tests, we displayed all the results using 99.4% confidence intervals instead of 95.0% intervals, which is a 5% type I error divided by eight tests.

In an unplanned analysis, we examined the association between the reviewer’s recommendation and whether they included citations to work other than citations to their own articles. This was added to examine differences between reviewers’ citations to their own articles and other articles.

Outliers were not excluded. No data were missing in the analysis data set.

The sample size calculation is in Supplement S.1.

Text analysis

We examined how reviewers’ citations to their own articles or other articles were justified and whether their wording differed according to their recommendation. For an initial view of citations to their own articles, we randomly selected 20 reviews and extracted the most relevant sentence concerning the citation.

To analyse the review text, we first extracted the 100 most commonly used words in all reviews. To standardise the text, all words were transformed into tokens, with stop-words removed and then stemmed. We then tested which of the 100 words were associated with recommending Approved versus Reservations or Not approved amongst those reviewers who included a citation to their own articles and those who included a citation to other articles. We chose the set of words using an elastic net with 10-fold cross-validation and selected a parsimonious model by using the lambda within one standard error of the minimum cross-validated error [38, 39]. To get uncertainty intervals for the estimates, we fitted a Bayesian model with the set of words selected by the elastic net and using a sceptical Normal prior centred on zero to create shrinkage.

Reproducibility

Research question 1 was pre-registered using As Predicted on 20 May 2024 [40]. Research question 2 was formulated during data collection but before any data analysis and used the same study design and statistical methods as question 1.

All data extraction and analyses were conducted using R version 4.4.1 [41]. The data and R code are available on GitHub [42].

Results

A flow chart of the included reviews is shown in Supplement S.2. The final sample size was over 37,000 reviews. There were more than 3,500 articles that were not included because they had not yet been peer reviewed, especially recent articles. More than 2,000 reviewers did not have a record in OpenAlex and so could not be included. These missing reviewers were more likely to be from older articles and more likely to be co-reviewers.

Descriptive statistics on the included reviews are in Table 2. The reviewers were cited at least once in 13% of the articles and 6% of the reviews included a self-citation. Most reviews recommended “Approved” (54%), with only 8% recommending “Not approved” which is low compared with many journals; however, 40 to 50% of submissions are rejected before articles are sent for peer review (personal communication, F1000 staff). The reviewers were relatively experienced, with a median number of papers of 55.

The binary predictor for citations of “any versus none” had a generally better fit to the data compared to the linear predictor (Supplement S.3). This indicates that for most reviewers receiving any citation is important, and there is no linear increase for two or more citations. The following results are for the binary predictor “any versus none”, with the results using a linear predictor in Supplement S.4.

Reviewers who were cited were more likely to approve the article, but only after version 1 (Fig 2 and Table 3). If a reviewer was cited in any versions after version 1, the odds ratio for recommending Approved versus Reservations or Not approved was 1.61 (adjusted 99.4% CI 1.16 to 2.23).

Odds ratios and probabilities for reviewers giving a more or less favourable recommendation depending on whether they were cited in the article.

Top left: Odds ratios for reviewers giving a more favourable (Approved) or less favourable (Reservations or Not approved) recommendation depending on whether they were cited in the article. Reviewers cited in later versions (blue) were more likely to make a favourable recommendation (odds ratio = 1.61; adjusted 99.4% CI: 1.16 to 2.23), whereas being cited in the first version (green) did not improve their recommendation (odds ratio = 0.84; adjusted 99.4% CI: 0.69 to 1.03). Top right: Same results as top left displayed as conditional probabilities. From the top, the lines show the within-strata probability of a reviewer approving: a version 1 article in which they are not cited (0.51; adjusted 99.4% CI: 0.49 to 0.52); a version 1 article in which they are cited (0.46; adjusted 99.4% CI: 0.41 to 0.51); a version 2 (or higher) article in which they are not cited (0.48; adjusted 99.4% CI: 0.45 to 0.51); and a version 2 (or higher) article in which they are cited (0.60; adjusted 99.4% CI: 0.53 to 0.65). Bottom left: Same estimates as top left except that a more favourable recommendation is now Approved or Reservations and a less favourable is Not approved. There was no clear association for cited reviewers in version 1 (odds ratio = 0.84; adjusted 99.4% CI: 0.57 to 1.23) or later versions (odds ratio = 1.12; adjusted 99.4% CI: 0.59 to 2.13). Bottom right: Same results as bottom left displayed as conditional probabilities. From the top, the lines show the within-strata probability of a reviewer approving: a version 1 article in which they are not cited (0.51; adjusted 99.4% CI: 0.48 to 0.53); a version 1 article in which they are cited (0.47; adjusted 99.4% CI: 0.35 to 0.55); a version 2 (or higher) article in which they are not cited (0.50; adjusted 99.4% CI: 0.43 to 0.55); and a version 2 (or higher) article in which they are cited (0.53; adjusted 99.4% CI: 0.34 to 0.63). This figure is based on an analysis of 12,051 articles and 24,677 reviews for version 1 and 6090 articles and 10,196 reviews for version 2+. In all panels a dot or square represents a mean, and a horizontal line represents an adjusted 99.4% confidence interval.

Reviewers who included a citation to their own articles were much less likely to approve the article for all versions (Fig 3 and Table 3). The odds ratio for recommending Approved versus Reservations or Not approved was 0.57 (99.4% CI 0.44 to 0.73) for version 1 and strengthened to 0.15 (99.4% CI 0.08 to 0.30) for versions 2+. The less favourable recommendation was only for the approval of the article and the odds ratios for Approved or Reservations versus Not approved were much closer to 1.

Odds ratios and probabilities for reviewers giving a more or less favourable recommendation if they included a citation to their own articles in their review.

Top left: Odds ratios for reviewers giving a more favourable (Approved) or less favourable (Reservations or Not approved) recommendation depending on whether their review included a citation to their own articles. Reviewers including a citation to their own articles were less likely to make a favourable recommendation for version 1 (green; odds ratio = 0.57; adjusted 99.4% CI: 0.44 to 0.73) and later versions (blue; odds ratio = 0.15; adjusted 99.4% CI: 0.07 to 0.30). Top right: Same results as top left displayed as conditional probabilities. From the top, the lines show the within-strata probability of a reviewer approving: a version 1 article in which their review did not include a citation (0.51; adjusted 99.4% CI: 0.50 to 0.53); a version 1 article in which in which their review included a citation (0.37; adjusted 99.4% CI: 0.29 to 0.44); a version 2 (or higher) article in which their review did not include a citation (0.51; adjusted 99.4% CI: 0.49 to 0.53); and a version 2 (or higher) article in which in which their review included a citation (0.14; adjusted 99.4% CI: 0.01 to 0.30). Bottom left: Same estimates as top left except that a more favourable recommendation is now Approved or Reservations and a less favourable is Not approved. There was no clear association for reviewers who included a citation to their own articles in version 1 (odds ratio = 1.11; adjusted 99.4% CI: 0.77 to 1.60) or later versions (odds ratio = 0.80; adjusted 99.4% CI: 0.37 to 1.74). Bottom right: Same results as bottom left displayed as conditional probabilities. From the top, the lines show the within-strata probability of a reviewer approving: a version 1 article in which their review did not include a citation (0.50; adjusted 99.4% CI: 0.48 to 0.52); a version 1 article in which their review included a citation (0.53; adjusted 99.4% CI: 0.43 to 0.60); a version 2 (or higher) article in which their review did not include a citation (0.51; adjusted 99.4% CI: 0.46 to 0.54); and a version 2 (or higher) article in which their included a citation (0.45; adjusted 99.4% CI: 0.12 to 0.61). This figure is based on an analysis of 12,078 articles and 24,732 reviews for version 1 and 6101 articles and 10,213 reviews for version 2+. In all panels a dot or square represents a mean, and a horizontal line represents an adjusted 99.4% confidence interval.

In an unplanned analysis, we examined the behaviour of reviewers in the first two versions of the article. We examined the 441 reviews where the reviewer was not cited in version 1 of the article and included a citation to their own articles in their first review. The reviewers who were then cited in version 2 recommended approval for 92% compared to 76% for reviewers who were not cited (odds ratio = 3.5, 95% CI: 2.0 to 6.1). This analysis did not use matching.

In an unplanned analysis, we examined whether the reviewers’ recommendations depended on whether their review included citations to articles other than their own. Reviewers who included citations in their review were much more likely not to approve the article (Figure 4), which was similar to the association with citations to reviewers’ own articles (Figure 3). However, reviewers who included citations to articles other than their own were also much more likely to recommend “Not approved”, as shown by the lower odds of “Approved” or “Reservations” versus “Not approved”. This association was not seen using citations to reviewers’ own articles (Figure 3)

Odds ratios and probabilities for reviewers giving a more or less favourable recommendation depending on if they included citations to articles other than their own in their review.

Top left: Odds ratios for reviewers giving a more favourable (Approved) or less favourable (Reservations or Not approved) recommendation depending on whether their review included a citation to articles other than their own. Reviewers including citations to other articles were less likely to make a favourable recommendation for version 1 (green; odds ratio = 0.53; adjusted 99.4% CI: 0.44 to 0.64) and later versions (blue; odds ratio = 0.18; adjusted 99.4% CI: 0.10 to 0.30). Top right: Same results as top left displayed as conditional probabilities. From the top, the lines show the within-strata probability of a reviewer approving: a version 1 article in which their review did not cite other articles (0.53; adjusted 99.4% CI: 0.51 to 0.54); a version 1 article in which their review cited other articles (0.37; adjusted 99.4% CI: 0.31 to 0.42); a version 2 (or higher) article in which their review did not cite other articles (0.52; adjusted 99.4% CI: 0.50 to 0.54); and a version 2 (or higher) article in which in which their review cited other articles (0.17; adjusted 99.4% CI: 0.02 to 0.30). Bottom left: Same estimates as top left except that a more favourable recommendation is now Approved or Reservations and a less favourable is Not approved. Reviewers including citations to other articles were less likely to make a favourable recommendation for version 1 (odds ratio = 0.62; adjusted 99.4% CI: 0.46 to 0.84) and later versions (odds ratio = 0.34; adjusted 99.4% CI: 0.16 to 0.73). Bottom right: Same results as bottom left displayed as conditional probabilities. From the top, the lines show the within-strata probability of a reviewer approving: a version 1 article in which their review did not cite other articles (0.52; adjusted 99.4% CI: 0.49 to 0.54); a version 1 article in which their review cited other articles (0.41; adjusted 99.4% CI: 0.31 to 0.48); a version 2 (or higher) article in which their review did not cite other articles (0.52; adjusted 99.4% CI: 0.47 to 0.55); and a version 2 (or higher) article in which their review cited other articles (0.27; adjusted 99.4% CI: 0.02 to 0.45). This figure is based on an analysis of 12,078 articles and 24,732 reviews for version 1 and 6101 articles and 10,213 reviews for version 2+. In all panels a dot or square represents a mean, and a horizontal line represents an adjusted 99.4% confidence interval.

Sensitivity analyses

The odds ratios when including co-reviewers with reviewers were similar to the odds ratios when using reviewers only (Supplement S.5).

We found no evidence that the reviewers’ publication numbers or country confounded the associations between citations and recommendations (Supplement S.6).

Text analyses of reviewers’ comments

Earlier version: A random sample of how reviewers included citations to their own articles found some vague justifications (Supplement S.7); for example, “Here are some additional publications you might consider referencing”. Other sentences adhered to the publisher’s guidelines for reviewers, as specific reasoning was provided for citations to their own articles [7]. One reviewer thanked the authors for a previous citation. Three reviews did not have a relevant sentence. One reviewer almost certainly used AI to write their review as it included the phrase “Certainly! Here are some potential review questions for the manuscript” [43]; this review included six self-citations with no justifications.

Revised version: A random sample of how reviewers included citations to their own articles found some vague justifications (Supplement S.7); for example, “Here are some additional publications you might consider referencing”. Other sentences adhered to the publisher’s guidelines for reviewers, as specific reasoning was provided for citations to their own articles [7]. One reviewer thanked the authors for a previous citation. Three reviews did not have a relevant sentence. One reviewer very likely used AI to write their review as it included the phrase “Certainly! Here are some potential review questions for the manuscript” [43]; this review included six self-citations with no justifications.

Reviewers who included a citation to their own articles or other articles were more likely to use the words “need” and “please” when not approving the article (Figure 5). In contrast, the words “genome” and “well” were the most strongly associated with the reviewers’ approval.

Words in the reviewers’ comments that were associated with approving the article or not for reviewers who included a citation to their own articles (n = 2,025) and reviewers who included citations to other articles (n = 4, 350).

The words were selected using an elastic net that started with the 100 most commonly used review words. The estimates from the elastic net are shown as empty circles and the mean estimates and 95% credible intervals from a Bayesian model are shown as a solid circle and horizontal line. Words are shown if the probability of a non-zero mean was over 0.95 for either reviewers who cited their own articles or reviewers who cited other articles. Four words were selected by the elastic net for the reviewers who cited other articles but not by the elastic net for reviewers who cited their own articles. The axis label shows the stemmed word and most common whole word in brackets.

To examine how often open peer reviews were viewed, we took a random sample of 200 reviews from the four journals and found that, on average, they were viewed just 1.2 times per year (Supplement S.8).

Discussion

Our results provide evidence that some reviewers have a transactional view of peer review, with their final approval dependent on citations to their work. Some reviewers may be exploiting the pressure on authors to “publish or perish”. Under this pressure, many authors may oblige and add the suggested citations, especially since adding another citation may only require a minor edit to their article [44]. Both sides gain from this transaction, as the authors get an indexed publication and the reviewer gets a citation.

A key question is whether citations to a reviewer’s own articles are justified as they may highlight important errors or missing context in the article. Citations to a reviewer’s own articles can be justified when the authors have made a “large scholarly oversight” [8]. To investigate this, we compared the recommendations and wording of reviewers who included citations to their own articles to reviewers who included citations to other articles. The language used by the reviewers of these two groups was similar, with a higher use of “please” and “need” when not approving the article (Figure 5). However, there was a difference between groups in their recommendations, as reviewers including citations to other articles were more likely to recommend “Not approved” (Figure 4) whereas this association was not observed for reviewers including citations to their own articles (Figure 3). This indicates that missing citations to other articles were considered more serious than missing citations to the reviewer’s articles. Reviewers who cited their own articles may have been more inclined to give authors a chance to update their article and thus potentially include the “missing” citation(s).

Examining the context of the citations to reviewers’ own articles, we found vague or non-existent justifications (Supplement S.7), showing that some reviewers ignored the journals’ guidelines to state their reasoning when including citations to their own articles. However, these examples of poor justifications do not mean that all self-citations are coercive.

For both research questions, the effects were stronger for the second and later versions of the article than for the first version. Reviewers may understand that authors may be more willing to compromise on later versions when they are closer to obtaining an indexed publication. Most researchers understand that the peer review system is imperfect and that they sometimes have to make compromises to be successful [20, 45]. Another difference to consider is that later versions will include more articles with disagreements between reviewers and more that were not “Approved” in the first version, as articles where two or more reviewers recommended “Approved” may not have needed a second version.

Potential improvements to peer review

Journals could give stronger guidance to reviewers and authors on coercive citations [4]. However, given the limited time for peer review and the many differences in guidelines between journals [46], most authors may not read peer review instructions. Hence, guidance alone may have limited impact.

One suggestion is that reviewers declare to editors when they have recommended citations to their own work [47]. A useful innovation would be for all reviews that contain citations to the reviewer’s own articles to be automatically flagged to the editors who could check if the citations are justified. We are aware of one journal where this is already happening (personal communication, Benno Torgler). F1000 have recently introduced checks to prevent reviewers from publishing a review with three or more citations to the reviewer’s own articles. If the reviewers continue to request more than three, then the review is examined, and if the citations are deemed inappropriate and the reviewer declines to remove them, then the review is declined.

Open peer review has been suggested as a way to reduce coercive citations [7, 47]. However, our results from four journals that use open review show that it is not a perfect antidote, although the problem could be worse in journals using blinded peer review. The transparency of open peer review should prevent reviewers from leaving self-serving comments; however, we found some dubious justifications for self-citations and blatant use of AI (Supplement S.7). These reviewers may have rationalised that although their words are public, they are rarely scrutinised (Supplement S.8); hence, it was worth the risk. The assumed additional quality assurance from open peer review [48] may often be absent.

A more radical change to peer review is that the reviewers initially see a version of the article with all references blinded and no reference list; for example, “A strong association between increased cleaning and reduced hospital infection is well established [x]”. Reviewers are asked to give an initial recommendation and comments, and then are shown the version with the full references and asked if they need to update their recommendation or provide additional comments. However, this involves more administrative work and demands more from peer reviewers. This approach could be used for particularly consequential or controversial articles. Some journals already require authors to partially blind their articles to maintain anonymous peer review; for example, the instructions from Taylor & Francis include blinding the authors’ names in the reference list [49].

An argument could be made for using large language models to provide peer review that is unmoved by citation flattery. However, peer review is an inherently human task by peers, and instead we need to improve peer review rather than abdicating this often difficult and time-consuming task to machines [50].

Related research

Previous cross-sectional studies of self-citations in reviews found at least one self-citation in 3% at a journal that used blinded peer review [17], 12% at a journal that used blinded peer review [51], and 12% at a journal that used open peer review [52]. A related study found that 15% of reviews included a self-citation and that the self-citations were highest when the reviewer recommended “major revisions” [53]. These figures are comparable with the 6% found here and indicate that most reviews do not include self-citations.

Previous surveys estimated that 14% and 20% of authors had experienced a coercive citation request from an editor [54, 55], and 7% and 23% had experienced coercive citation pressure from a reviewer [14, 56]. The frequency with which researchers interact with peer review means that many will encounter coercive citations at some point in their career.

A study of conference submissions estimated that reviewers who were cited gave submissions much higher scores [18]. A study of journal peer review estimated that cited reviewers scored the article higher, but with potential confounding by the quality of the article [17].

A survey of authors concluded that accepting an editor’s request for citations improved the chances of being accepted [12]. Requests in later versions were more strongly associated with acquiescence, and we found a related pattern in our analysis, with reviewers who included citations to their own articles being much less likely to recommend approval for later versions (Figure 3). A study examining open peer review found that requests to cite the reviewer’s articles were more likely to be included than other suggested citations, indicating that many authors wanted to please the reviewer or felt pressure to do so [52].

A survey of journal editors found that only 5% objected to reviewers citing their own articles, and that this should be expected as reviewers are likely to have done related work [8].

A cross-sectional study found that reviewers citations to their own articles were more likely to have no rationale compared to other citations, suggesting that they are more likely to be unwarranted [51].

Strengths and limitations

This is an observational study meaning we cannot rule out unmeasured confounding and should be cautious in interpreting the results.

To our knowledge, this is the first analysis to use a matched design and analysis when examining reviewer citations, and hence strongly control for any confounding by the characteristics of the authors or articles. We compared reviewers who examined an identical article; hence, the differences we found should be due to the reviewers.

Our models include measurement error, as some citations to the reviewers’ work will be missed by our data collection, and some captured citations will be inaccurate [57]. We performed random data checks that showed good accuracy (Supplement S.9); however, we also found valid citations that were not captured by our data extraction for conference proceedings and technical reports, which are less likely to have a DOI. This measurement error would most likely underestimate a true association, as it reduced the variance in citation counts and created a regression dilution [58]. Our estimates will be biased if the associations between citations and reviewers’ recommendations are different for publications that do not have a DOI. Reviewers should be equally happy with any citation to their work; however, some reviewers may prefer citations to indexed articles, as these are more likely to count toward their h-indices [59].

We examined whether citing a reviewer altered their recommendation, but did not examine the sentiment of the citation [60]. Some citations would likely have been critical of the reviewer’s articles, and we would expect these to reduce the chances of a favourable recommendation. An analysis that included the sentiment of the citation would be useful, although previous research found that most citations are neutral or positive [60].

We did not examine the authors’ responses to the reviewers but these could include important information on why a citation was included or not in a revised version of the article. A detailed analysis examining the text used in the interactions between authors and reviewers could provide valuable information about the peer review process.

Our results may not be generalisable to journals that use blinded peer review or journals that use the traditional peer review model rather than the publish-review-curate model studied here. A previous study found that asking reviewers to consent to an open review had no important effect on the quality of the review or the reviewers’ recommendation [61]. Another potential difference is that the journals in our sample often asked the authors to suggest peer reviewers; however, this is relatively common in other journals [8].

We found a bias in our sample, as co-reviewers and reviewers from older articles were more likely to be excluded due to not having an OpenAlex record (Supplement S.6). We therefore lost more junior reviewers who were less likely to be cited. The percentage of reviews lost was 5% (2,026 of 39,113), which is hopefully small enough to avoid a large bias.